---
title: "Distributions"
author: "Guyliann Engels & Philippe Grosjean"
description: "**SDD I Module 7** Distributions binomiale et Poisson."
tutorial:
  id: "A07Lb_distri"
  version: 2.0.0/9
output:
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
BioDataScience1::learnr_setup()
SciViews::R()
```

```{r, echo=FALSE}
BioDataScience1::learnr_banner()
```

```{r, context="server"}
BioDataScience1::learnr_server(input, output, session)
```

----

## Objectifs

Les lois de distribution généralisent le calcul des probabilités dans des situations bien définies. Elles permettent de calculer la probabilité que des évènements se produisent d'un point de vue théorique. La distribution binomiale et celle de Poisson concernent des variables quantitatives à deux modalités, c'est-à-dire que seulement deux évènements disjoints peuvent se produire.

- Vérifiez vos connaissances relatives à la loi de distribution des probabilités binomiale

- Vous assurez d'avoir bien compris la distribution de Poisson

Vous devez avoir étudié le contenu du [module 7](https://wp.sciviews.org/sdd-umons/?iframe=wp.sciviews.org/sdd-umons-2020/proba.html) du cours, et en particulier les sections relatives à la [distribution binomiale](https://wp.sciviews.org/sdd-umons/?iframe=wp.sciviews.org/sdd-umons-2020/distribution-binomiale.html) et à la [distribution de Poisson](https://wp.sciviews.org/sdd-umons/?iframe=wp.sciviews.org/sdd-umons-2020/distribution-de-poisson.html). Enfin, cette matière nécessite que vous soyez à l'aise avec le calcul des probabilités, et que vous l'ayez vérifié via le learnr `BioDataScience1::run("A07La_proba")`.

## Distribution binomiale

La distribution binomiale calcule les probabilités de *n* succès pour un modèle essai-erreur indépendant. Cela signifie que seulement deux évènements disjoints sont possibles\ : le "succès" ou l'"échec" (ces noms sont attribués aux deux évènements qui peuvent représenter tout autre chose comme pile ou face, mâle ou femelle, chevelu ou non, ...). Le nombre d'essais indépendants *p* est fixé à l'avance et est toujours le même pour une distribution donnée.

![](images/lions.jpg)

En guise d'entrée en matière, répondez aux questions suivantes en effectuant le calcul à la main (en vous aidant éventuellement d'une calculette). Considérant une population de lions dans une réserve naturelle au Kenya qui est composée d'autant de lions mâles que de lionnes (sexe ratio de 1:1). Vous vous baladez dans la réserve et vous avez observé dix individus au total.

```{r qu_binom}
quiz(
  question("Calculez la probabilité d'avoir observé trois mâles au maximum.",
    answer("0.117"),
    answer("0.161"), 
    answer("0.172", correct = TRUE),
    answer("0.828"),
    allow_retry = TRUE, random_answer_order = TRUE,
    incorrect = "Recommencez afin de trouver la bonne réponse",
    correct = "C'est correct ! Si l'on considère une distribution aléatoire des individus dans la réserve, du moins, pour cet exemple fictif."),
  question("Calculez la probabilité d'observer au minimum quatre mâles parmi ces dix individus.",
    answer("0.117"),
    answer("0.161"), 
    answer("0.172"),
    answer("0.828", correct = TRUE),
    allow_retry = TRUE, random_answer_order = TRUE,
    incorrect = "Recommencez afin de trouver la bonne réponse",
    correct = "C'est correct ! Notez que cette probabilité est complémentaire à la précédente.")
)
```

Calculez la table des probabilités possibles pour tous les évènements depuis zéro jusqu'à dix lions mâles observés sur dix individus au total, avec R.

💬 **Ce code correspond au snippet `.ibtable`** [`.ib`= (d)istribution: `b`inomial].

```{r binom1_h2, exercise=TRUE}
(.table <- data.frame(success = ___,
    probability = dbinom(___, size = ___, prob = ___)))
```

```{r binom1_h2-hint-1}
(.table <- data.frame(success = 0:10,
    probability = dbinom(0:10, size = ___, prob = ___)))

 #### ATTENTION: Hint suivant = solution !####
```

```{r binom1_h2-solution}
(.table <- data.frame(success = 0:10,
    probability = dbinom(0:10, size = 10, prob = 0.5)))
```

```{r binom1_h2-check}
grade_code("C'est parfait.")
```

Employez R cette fois-ci pour répondre à une des questions posées lors du quiz plus haut\ : "Calculez la probabilité d'observer au maximum 3 mâles sur 10 individus."

💬 **Ce code correspond au snippet `.ibproba`.**

```{r binom2_h3, exercise=TRUE}
pbinom(___, size = ___, prob = ___, lower.tail = ___)
```

```{r binom2_h3-hint-1}
pbinom(___, size = ___, prob = ___, lower.tail = TRUE)
```

```{r binom2_h3-hint-2}
pbinom(3, size = ___, prob = ___, lower.tail = TRUE)

 #### ATTENTION: Hint suivant = solution !####
```

```{r binom2_h3-solution}
pbinom(3, size = 10, prob = 0.5, lower.tail = TRUE)
```

```{r binom2_h3-check}
grade_code("Vous comprenez manifestement bien la logique de ce type de calcul dans R... ou alors, vous avez passé énormément de temps à observer les lions au Kenya !")
```

Calculez avec R la probabilité d’observer exactement 4 lionnes pour 10 contacts dans une autre réserve où il y a 3 femelles pour 1 mâle (Pr{femelle} = 0.75).

```{r binom3_h3, exercise=TRUE}
pbinom(___, size = ___, prob = ___, lower.tail = ___) - ___
```

```{r binom3_h3-hint-1}
pbinom(___, size = ___, prob = ___, lower.tail = ___) - pbinom(___, size = ___, prob = ___, lower.tail = ___)
```

```{r binom3_h3-hint-2}
pbinom(4, size = ___, prob = ___, lower.tail = TRUE) - pbinom(___, size = ___, prob = ___, lower.tail = TRUE)

 #### ATTENTION: Hint suivant = solution !####
```

```{r binom3_h3-solution}
pbinom(4, size = 10, prob = 0.75, lower.tail = TRUE) - pbinom(3, size = 10, prob = 0.75, lower.tail = TRUE)
```

```{r binom3_h3-check}
grade_code("Comme le calcul des probabilités selon des lois de distribution nous renvoie les probabilités d'être plus petit ou égal avec `lower.tail = TRUE` ou plus grand ou égale avec `lower.tail = FALSE`, nous devons ruser et soustraire deux probabilités pour avoir un intervalle borné des deux côtés (ou ici, une seule valeur possible).")
```

Représentez le graphique de densité de probabilité pour la distribution binomiale de l'exercice précédent.

💬 **Ce code correspond au snippet `.ibdens`.**

```{r binom4_h2, exercise=TRUE}
plot(___, dbinom(___, size = ___, prob = ___), type = "h",
  col = "black", xlab = "Quantiles", ylab = "Probability mass")
```

```{r binom4_h2-hint-1}
plot(0:10, dbinom(0:10, size = ___, prob = ___), type = "h",
  col = "black", xlab = "Quantiles", ylab = "Probability mass")

 #### ATTENTION: Hint suivant = solution !####
```

```{r binom4_h2-solution}
plot(0:10, dbinom(0:10, size = 10, prob = 0.75), type = "h",
  col = "black", xlab = "Quantiles", ylab = "Probability mass")
```

```{r binom4_h2-check}
grade_code("Les barres verticales représentent ici la probabilité de chaque évènement discret correspondant au nombre de succès (les quantiles) en fonction du nombre d'essais indépendants qui est toujours le même (indiqué via `size=`).")
```

## Distribution de Poisson

![](images/fishing.jpg)

La distribution de Poisson détermine, tout comme la distribution binomiale, les probabilités de nombres de succès par rapport à un nombre d'essais indépendants, mais ici la probabilité de succès est très faible (évènement rare). Vous pouvez utiliser cette distribution pour déterminer la probabilité de rencontrer un animal rare, pour prédire les tremblements de terre, pour déterminer la probabilité d'une maladie rare, ... et même pour prédire la probabilité que vous arriviez à pêcher un gros poisson avec votre petite canne à pêche\ !

Répondez à la question suivante en effectuant le calcul à la main (en vous aidant éventuellement d'une calculette).

```{r qu_poisson}
quiz(
  question("Calculez le troisième terme d'une distribution de poisson dont le $\\lambda = 3$.",
    answer("0.353"),
    answer("0.647"), 
    answer(" 0.224", correct = TRUE),
    answer("0.828"),
    allow_retry = TRUE, random_answer_order = TRUE,
    incorrect = "Recommencez afin de trouver la bonne réponse",
    correct = "C'est correct !")
)
```

Utilisez R pour représentez la table de probabilités lié à l'exercice ci-dessus.

💬 **Ce code correspond au snippet `.iptable`** [`.ip` = (d)`i`stribution: `p`oisson].

```{r poisson1_h2, exercise=TRUE}
(.table <- data.frame(occurences = 0:(___+20), probability = dpois(0:(___+20),
  lambda = ___)))
```

```{r poisson1_h2-hint-1}
(.table <- data.frame(occurences = 0:(3+20), probability = dpois(0:(___+20),
  lambda = ___)))

 #### ATTENTION: Hint suivant = solution !####
```

```{r poisson1_h2-solution}
(.table <- data.frame(occurences = 0:(3+20), probability = dpois(0:(3+20),
  lambda = 3)))
```

```{r poisson1_h2-check}
grade_code("La distribution de Poisson est plus simple que la distribution binomiale car elle n'admet qu'un seul paramètre, lambda, correspondant au nombre de succès le plus probable.")
```

Représentez le graphique de densité de la distribution de Poisson pour $\lambda = 12$.

💬 **Ce code correspond au snippet `.ipdens`.**

```{r poisson2_h2, exercise = TRUE}
plot(0:(___+20), dpois(0:(___+20), lambda = ___),
  type = "h", col = "black", xlab = "Quantiles", ylab = "Probability mass")
```

```{r poisson2_h2-hint-1}
plot(0:(12+20), dpois(0:(___+20), lambda = ___),
  type = "h", col = "black", xlab = "Quantiles", ylab = "Probability mass")

 #### ATTENTION: Hint suivant = solution !####
```

```{r poisson2_h2-solution}
plot(0:(12+20), dpois(0:(12+20), lambda = 12),
  type = "h", col = "black", xlab = "Quantiles", ylab = "Probability mass")
```

```{r poisson2_h2-check}
grade_code("Ce graphique est en fait très asymétrique, mais l'axe des quantiles est tronqué vers la droite puisque toutes les évènements avec un plus grand nombre de succès ont tous une probabilité très proche de zéro et ne sont donc pas représentés sur le graphique.")
```

## Conclusion

Bravo\ ! Vous venez de terminer votre auto-évaluation relative à deux des lois de distribution de probabilités les plus courantes. La logique et la façon d'effectuer des calculs de probabilités à partir de quantiles ou de quantiles à partir de probabilités est un exercice de base en statistiques que vous serez amené à utiliser fréquemment. Assurez-vous donc de bien la maîtriser.

```{r comm_noscore, echo=FALSE}
question_text(
  "Laissez-nous vos impressions sur cet outil pédagogique",
  answer("", TRUE, message = "Pas de commentaires... C'est bien aussi."),
  incorrect = "Vos commentaires sont enregistrés.",
  placeholder = "Entrez vos commentaires ici...",
  allow_retry = TRUE
)
```
