---
title: "Test t de Student"
author: "Guyliann Engels & Philippe Grosjean"
description: "**SDD I Module 9** La moyenne et le test t de Student."
tutorial:
  id: "A09La_ttest"
  version: 2.0.0/9
output:
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
BioDataScience1::learnr_setup()
SciViews::R()
```

```{r, echo=FALSE}
BioDataScience1::learnr_banner()
```

```{r, context="server"}
BioDataScience1::learnr_server(input, output, session)
```

----

## Objectifs

La moyenne est un descripteur très employé pour résumer l'information. Il est très courant de comparer un moyenne à une valeurs de référence ou deux moyennes ensembles. Dans votre carrière scientifique, il est indispensable de maîtriser la moyenne et ses subtilités.

- Savoir calculer la moyenne, l'écart-type ou encore la variance

- Comprendre le test *t* de Student et être capable de l'utiliser pour résoudre des questions pratiques en biologie

## Moyenne

Commençons par revoir les bases. Un échantillon représentatif d'une population peut être décrit par plusieurs descripteurs statistiques comme la moyenne, l'écart-type ou encore le coefficient de variation.

### Moyenne du vecteur `vec`

Utilisez une fonction dans R afin de calculer la moyenne du vecteur `vec`.

```{r, echo = TRUE}
vec <- c(14, 24, 32, 18, 19)
```

```{r mean1_h2, exercise = TRUE}
(vec <- c(14, 24, 32, 18, 19))
# Calculez ci-dessous la moyenne de vec

```

```{r mean1_h2-hint-1}
(vec <- c(14, 24, 32, 18, 19))
# Calculez ci-dessous la moyenne de vec : 
# Utilisez la foncion mean()

 #### ATTENTION: Hint suivant = solution !####
```

```{r mean1_h2-solution}
(vec <- c(14, 24, 32, 18, 19))
## Solution ##
mean(vec)
```

```{r mean1_h2-check}
grade_result(
  pass_if(~ identical(.result, mean(vec)), "Bien joué ! La fonction mean() est la bonne fonction."),
  fail_if(~ TRUE, "Votre calcul n'est pas le bon. Retentez votre chance. Réfléchissez à la bonne fonction à employer.")
)
```

### Moyenne du vecteur `vec2`

Utilisez une fonction dans R afin de calculer la moyenne du vecteur `vec2`.

```{r, echo = TRUE}
vec2 <- c(14, 3, 32, NA, 40)
```

```{r mean2_h3, exercise = TRUE}
(vec2 <- c(14, 3, 32, NA, 40))
# Calculez ci-dessous la moyenne de vec

```

```{r mean2_h3-hint-1}
(vec2 <- c(14, 3, 32, NA, 40))
# Calculez ci-dessous la moyenne de vec : 
# Utilisez la foncion mean()

 #### ATTENTION: Hint suivant = solution !####
```

```{r mean2_h3-hint-2}
(vec2 <- c(14, 3, 32, NA, 40))
# Calculez ci-dessous la moyenne de vec : 
# Utilisez la foncion mean(), utilisez l'argument na.rm = 

```

```{r mean2_h3-solution}
(vec2 <- c(14, 3, 32, NA, 40))
## Solution ##
mean(vec, na.rm = TRUE)
```

```{r mean2_h3-check}
grade_result(
  pass_if(~ identical(.result, mean(vec2, na.rm = TRUE)), "Bien joué ! La fonction mean() est la bonne idée. La petite suptilité est sur l'utilisation de l'argument na.rm = TRUE."),
  fail_if(~ TRUE, "La bonne idée est d'employer la fonction mean(). Il y a une astuce avec les arguments à employer.")
)
```

### Ecart-type

Un second descripteur d'un échantillon est l'écart-type qui est calculé via l'équation suivante :

$$s = \sqrt {s^2}$$

L'écart-type est la racine carré de la variance qui est également un descripteur statistique de l'échantillon. 

La variance de l'échantillon qui est calculée via l'équation suivante :

$$s^2 = \sum\limits_{i=1}^n \frac{(y_i - \bar y)^2}{n-1}$$

Il existe également des fonctions dans R afin de calculer la variance et l'écart-type. Calculez l'écart-type du vecteur `vec3`

```{r, echo = TRUE}
vec3 <- c(22, 18, 4, 24, 34, 3, 9)
```

```{r sd1_h2, exercise = TRUE}
vec3 <- c(22, 18, 4, 24, 34, 3, 9)
# Calculez l'écart-type

```

```{r sd1_h2-hint-1}
vec3 <- c(22, 18, 4, 24, 34, 3, 9)
# Calculez l'écart-type utilisez la fonction sd()

 #### ATTENTION: Hint suivant = solution !####
```

```{r sd1_h2-solution}
vec3 <- c(22, 18, 4, 24, 34, 3, 9)
# Calculez l'écart-type
sd(vec3)
```

```{r sd1_h2-check}
grade_result(
  pass_if(~ identical(.result, sd(vec3)), "Bien joué ! La fonction sd() est la fonction pour calculer l'écart-type."),
  fail_if(~ TRUE, "Votre calcul n'est pas le bon. Retentez votre chance. Il semble que vous n'avez pas utilisé la bonne fonction.")
)
```


### Coefficient de variation 

Un autre descripteur est le coefficient de variation de l'échantillon qui est calculé via 

$$cv \ (\%) = \frac{s}{\bar y} \times 100$$

Calculez le coefficient de variation du vecteur `vec3`

```{r, echo = TRUE}
vec3 <- c(22, 18, 4, 24, 34, 3, 9)
```

```{r cv1_h2, exercise = TRUE}
vec3 <- c(22, 18, 4, 24, 34, 3, 9)
# Calculez le coefficient de variation 

```

```{r cv1_h2-hint-1}
vec3 <- c(22, 18, 4, 24, 34, 3, 9)
# Calculez le coefficient de variation : utilisez la fonction sd() et la fonction mean

 #### ATTENTION: Hint suivant = solution !####
```

```{r cv1_h2-solution}
vec3 <- c(22, 18, 4, 24, 34, 3, 9)
# Calculez l'écart-type
sd(vec3)/mean(vec3)*100
```

```{r cv1_h2-check}
grade_result(
  pass_if(~ identical(.result, (sd(vec3)/mean(vec3))*100), "Bien joué ! Vous avez judicieusement utilisé la moyenne et l'écart-type pour obtenir le coefficient de variation."),
  fail_if(~ TRUE, "Votre calcul n'est pas le bon. Commencez par déterminer la moyenne et l'écart-type de ce vecteur. N'oubliez pas de multiplier votre réponse par 100.")
)
```

Ces premiers exercices vous ont permis de revoir les bases théoriques de l'utilisation de R et du calcul de la moyenne.

## Indice de Masse Corporelle

Il est temps de passer à du concret. Vous allez vous intéresser à l'indice de masse corporelle dans une population d'adulte ayant strictement plus de 18 ans et de moins de 30 ans. Ce tableau de données ne vous est pas inconnu. Vous l'avez déjà employé à plusieurs reprises pour réaliser des graphiques.

![](images/diet.jpg){width='60%'}


```{r, echo=TRUE}
biometry <- read("biometry", package = "BioDataScience", lang = "fr")
glimpse(biometry)
```

Débutez par sélectionner les individus ayant strictement plus de 18 ans et 30 ans ou moins.

```{r bio_filt_h2, exercise = TRUE}
# Importation des données
biometry <- read("biometry", package = "BioDataScience") 
# Sélection des données
biometry <- ___(biometry, age ___ 18 ___ age ___ 30) 
# Première ligne du tableau 
head(biometry)
```

```{r bio_filt_h2-hint-1}
# Importation des données
biometry <- read("biometry", package = "BioDataScience") 
# Sélection des données
biometry <- filter(biometry, age > 18 ___ age ___ 30) 
# Première ligne du tableau 
head(biometry)
```

```{r bio_filt_h2-solution}
# Importation des données
biometry <- read("biometry", package = "BioDataScience") 
# Sélection des données
biometry <- filter(biometry, age > 18 & age <= 30) 
# Première ligne du tableau 
head(biometry)
```

```{r bio_filt_h2-check}
grade_code("Bien joué ! Vous avez proposé un filtre de votre tableau de données initiales.", "Oups, il semble que vous avez mal utilisé vos conditions.")
```

Vos supérieurs souhaitent savoir si il y a une différence significative de l'IMC entre les hommes et les femmes. La variable `bmi` est calculée pour vous.

```{r, echo = TRUE}
biometry %>.%
  mutate(., bmi = (weight/(height/100)^2)) %>.%
  labelise(., label = list(bmi = "IMC"), units = list(bmi = "")) -> biometry
```

Proposez un tableau résumé proposant la moyenne, l'écart-type et le nombre d'observations pour les hommes et pour les femmes.

```{r bio-prepare}
biometry <- read("biometry", package = "BioDataScience") %>.%
  filter(., age > 18 & age <= 30) %>.%
  mutate(., bmi = (weight/(height/100)^2)) %>.%
  labelise(., label = list(bmi = "IMC"), units = list(bmi = ""))
```

```{r bio_tab_h2, exercise = TRUE, exercise.setup = "bio-prepare"}
biometry %>.%
  ___(., ___) %>.%
  summarise(., mean = ___(___), sd = ___(___), n = n())
```

```{r bio_tab_h2-hint-1}
biometry %>.%
  group_by(., gender) %>.%
  summarise(., mean = ___(___), sd = ___(___), n = n())
#### ATTENTION: Hint suivant = solution !####
```

```{r bio_tab_h2-solution}
biometry %>.%
  group_by(., gender) %>.%
  summarise(., mean = mean(bmi), sd = sd(bmi), n = n())
```

```{r bio_tab_h2-check}
grade_code("Bien joué ! Vous avez proposé un tableau résumant l'information. Vous avez calculé la moyenne et l'écart-type de l'IMC en fonction du genre.", "Analysez avec attention le tableau que vous devez obtenir. Quelle est la variable dont vous devez déterminer la moyenne et l'écart-type ?")
```

Vous devez à présent comparer si l'indice de masse corporelle entre les hommes et les femmes. Utilisez un test t de Student bilatéral et de variance similaire. Vous devez fixer votre seuil $\alpha$ à 0.05. 

💬 **Un snippet peut vous aider à réaliser cet exercice `.hmtestindep`.**

```{r bio_ttest_h2, exercise = TRUE, exercise.setup = "bio-prepare"}
t.test(data = ___, ___ ~ ___,
  alternative = ___, conf.level = ___, var.equal = TRUE)
```

```{r bio_ttest_h2-hint-1}
t.test(data = ___, ___ ~ ___,
  alternative = "two.sided", conf.level = 0.95, var.equal = TRUE)
#### ATTENTION: Hint suivant = solution !####
```

```{r bio_ttest_h2-solution}
t.test(data = biometry, bmi ~ gender,
  alternative = "two.sided", conf.level = 0.95, var.equal = TRUE)
```

```{r bio_ttest_h2-check}
grade_code("Bien joué ! Vous avez su compléter la fonction t.test().", "Ce n'est pas la bonne réponse. Relis avec attention l'énoncé et complète les éléments manquants dans cette fonction.")
```

```{r ttest_quiz}
quiz(
  question("Est ce que l'indice de masse corporelle est similaire entre les hommes et les femmes ?",
           answer("oui", correct = TRUE),
           answer("non"),
           allow_retry = TRUE,
           incorrect = "Ce n'est pas la bonne réponse. Vous devez comparer la valeur de p au seuil $\alpha$  défini avant de réaliser le test pour trouver la bonne réponse à cette question.",
           correct = "Bravo, c'est correct. Vous avez bien interprété votre test."),
  question("Rejettez vous $H_0$ ?",
           answer("oui"),
           answer("non", correct = TRUE),
           allow_retry = TRUE,
           incorrect = "Vous vous êtes trompé. Pour savoir si vous devez rejeter $H_0*, il faut comparer la valeur de p au seuil $\alpha$  défini avant de réaliser le test.",
           correct = "Bravo, c'est correct. Vous ne rejettez pas le $H_0$ au seuil $\alpha$ de 0.05")
)
```


## Conclusion

Bravo! Vous venez de terminer votre séance d'exercices dans un tutoriel "learnr". 

```{r comm_noscore, echo=FALSE}
question_text(
  "Laissez-nous vos impressions sur cet outil pédagogique",
  answer("", TRUE, message = "Pas de commentaires... C'est bien aussi."),
  incorrect = "Vos commentaires sont enregistrés.",
  placeholder = "Entrez vos commentaires ici...",
  allow_retry = TRUE
)
```

