---
title: "Calcul de probabilités"
author: "Guyliann Engels & Philippe Grosjean"
description: "**SDD I Module 7** Calculer des probabilités dans R."
tutorial:
  id: "A07La_proba"
  version: 2.0.0/10
output:
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
BioDataScience1::learnr_setup()
SciViews::R()
```

```{r, echo=FALSE}
BioDataScience1::learnr_banner()
```

```{r, context="server"}
BioDataScience1::learnr_server(input, output, session)
```

------------------------------------------------------------------------

<!-- TODO: il manque des exercices qui comparent population de taille infinie ou non et qui mettent en lumière les différences au niveau du calcul des probabilités (conditionnelles) -->

## Objectifs

Le calcul des probabilités et le B-A-BA des statistiques. Il n'est pas toujours très intuitif, aussi, vous devez vous exercez et bien en comprendre les subtilités. Ce tutoriel vous permet de :

-   Appréhender le calculs de probabilités

-   Calculer des probabilités sur base d'un tableau de contingence

Avant d'aborder ce tutoriel, assurez-vous d'avoir bien compris le contenu de la [section 7.1](https://wp.sciviews.org/sdd-umons/?iframe=wp.sciviews.org/sdd-umons-2020/probabilit%25C3%25A9s.html) du cours. En effet, ce learnr sert d'auto-évaluation sur cette matière et ne sera utile que dans un contexte où vous la maîtrisez déjà, à des fins de vérification de vos acquis.

## Lire des probabilités

![](images/formalism.jpg)

Le calcul des probabilité est assorti d'un formalisme mathématique qu'il vous faut maîtriser. Les questions suivantes vous permettent de vérifier que vous le comprenez bien.

```{r qu_formalism}
quiz(
  question("Comment lire la proposition suivante : Pr{A.B} ?",
    answer("Probabilité d'observer A ou B"),
    answer("Probabilité d'observer A et B", correct = TRUE, message = "Probabilité d'évènements successifs indépendants qui se produisent simultanément"),
    answer("Probabilité d'observer A si B se produit"),
    allow_retry = TRUE, random_answer_order = TRUE,
    incorrect = "Pas exactement. Recommencez afin de trouver la bonne réponse.",
    correct = "C'est correct !"),
  question("Comment lire la proposition suivante : Pr{A|B} ?",
    answer("Probabilité d'observer A ou B"),
    answer("Probabilité d'observer A et B"),
    answer("Probabilité d'observer A si B se produit", correct = TRUE, message = "Ici, nous avons une probabilité conditionnelle que A se produise, sachant que B s'est produit"),
    allow_retry = TRUE, random_answer_order = TRUE,
    incorrect = "Désolé, mais ce n'est pas correct. Recommencez afin de trouver la bonne réponse",
    correct = "C'est correct !"),
  question("Comment lire la proposition suivante : Pr{A + B} ?",
    answer("Probabilité d'avoir A ou B", correct = TRUE, message = "En calcul de probabilité, l'addition signifie que l'on considère la probabilité que l'un parmi plusieurs évènements disjoints se produise"),
    answer("Probabilité d'avoir A et B"),
    answer("Probabilité d'avoir A si B se produit"),
    allow_retry = TRUE, random_answer_order = TRUE,
    incorrect = "Concentrez-vous sur ce que l'opérateur '+' signifie en calcul des probabilités. Recommencez afin de trouver la bonne réponse",
    correct = "C'est correct !"),
  question("Comment écrira-t-on la probabilité d'obtenir quatre mâles Pr{M} dans une portée de quatre châtons ?",
    answer("Pr{M . M . M . M}", correct = TRUE),
    answer("Pr{M et M et M et M}", correct = TRUE),
    answer("Pr{M + M + M + M}" ),
    answer("Pr{M ou M ou M ou M}"),
    allow_retry = TRUE, random_answer_order = TRUE,
    incorrect = "Est-ce des évènements successifs indépendants ou des évènements dépendants, réfléchissez. Recommencez afin de trouver la bonne réponse",
    correct = "C'est correct ! Le '.' est synonyme ici de 'et' et concerne des évènements indépendants successifs.")
)
```

## Calculer des probabilités

![](images/calculation.jpg)

### Stress

![](images/stress.jpg)

Le résultat d'une enquête relative au stress dans une population active en fonction de leurs revenus professionnels est le suivant :

```{r}
stress <- dtf(
  `Revenus faibles` = c( 526, 1954, 2480),
  `Revenus moyens`  = c( 274, 1680, 1954),
  `Revenus élevés`  = c( 216, 1899, 2115),
  Total             = c(1016, 5533, 6549))
rownames(stress) <- c("Stressé", "Non stressé", "Total")
knitr::kable(stress)
```

Calculez les probabilités suivantes relatives à cette étude (vous pouvez vous aider d'une calculette) :

```{r qu_calcul}
quiz(caption = "Stress et revenus",
  question("Quelle est la probabilité d'être stressé ?",
    answer("0.15", correct = TRUE , message = "Résultat issu du calcul suivant : 1016/6549, total stressés sur total général"),
    answer("0.52"),
    answer(" 0.08"),
    answer(" 0.30"),
    allow_retry = TRUE, random_answer_order = TRUE,
    incorrect = "Vous avez fait une erreur. Recommencez afin de trouver la bonne réponse",
    correct = "C'est correct !"),
  question("Quelle est la probabilité d'être stressé et d'avoir des revenus moyens ?",
    answer("0.042", correct = TRUE, message =  "Cette réponse est obtenue comme suit : 274/6549, nombre de stressés à revenus moyens sur le total général"),
    answer("0.52"),
    answer("0.15"),
    answer("0.08"),
    allow_retry = TRUE, random_answer_order = TRUE,
    incorrect = "Concentrez-vous et recommencez afin de trouver la bonne réponse",
    correct = "C'est correct !"),
  question("Pr{revenus moyens ou stressé} ?",
    answer("0.41", correct = TRUE, message =  "Cette réponse s'obtient en effet comme ceci : (1954/6549) + (1016/6549) - (274/6549), probabilité d'être stressé + probabilité d'avoir des revenus moyens. Mais les stressés à revenus moyens sont ici comptabilisés deux fois... donc on doit les soustraire un fois du calcul"),
    answer("0.042"),
    answer("0.52"),
    answer("0.086"),
    allow_retry = TRUE, random_answer_order = TRUE,
    incorrect = "Pas exactement. Recommencez afin de trouver la bonne réponse",
    correct = "C'est correct !"),
  question("Pr{revenus faibles | non stressé}",
    answer("0.35", correct = TRUE, message = "Le calcul est le suivant : 1954/5533, ici, nous considérons uniquement la population de non stressés. Donc on divise les revenus faibles par le total de non stressés et pas par le total général."),
    answer("0.41"),
    answer("0.52"),
    answer("0.15"),
    allow_retry = TRUE, random_answer_order = TRUE,
    incorrect = "C'est une probabilité conditionnelle. Elle se traite d'une façon particulière dans un tableau de contingence. Recommencez afin de trouver la bonne réponse",
    correct = "C'est correct !")
)
```

### Efficacité d'un médicament

![](images/pills.jpg)

Un nouveau médicament contre une maladie permet de soigner 80% des individus atteints. Malheureusement, une personne sur deux ne pourra pas prendre ce médicament à cause d'effets secondaires (patients « sensibles »). Considérant un nouveau malade arrivant à l'hôpital, quelle est la probabilité qu'il ne puisse pas être soigné par ce traitement ?

```{r medic1_h3, exercise=TRUE}

```

```{r medic1_h3-hint-1}
# Vous pouvez raisonner à partir de 100 patients
# Combien sont sensibles aux effets secondaires et ne pourront pas être soignés ?
```

```{r medic1_h3-hint-2}
# Vous pouvez raisonner à partir de 100 patients
# Combien sont sensibles aux effets secondaires et ne pourront pas être soignés ?
# Parmi ceux pouvant être soignés, combien ne guériront pas ?
# Additionnez les deux... et divisez par 100 !

 #### ATTENTION: Hint suivant = solution !####
```

```{r medic1_h3-solution}
# Patients ne pouvant pas recevoir le médicament + ceux ne guérissant pas
# avec ce médicament
0.5 + (0.2 * 0.5)
```

```{r medic1_h3-check}
grade_result(
  pass_if(~ identical(.result, 0.5 + (0.2 * 0.5)), "Il semble que le calcul des probabilités n'a plus de secrets pour vous. Nous avons deux sous-populations indépendantes : celle qui peut recevoir le médicament et celle qui ne le peut pas (nous additionons donc les probabilités respectives). Parmi ceux qui reçoivent le médicament, 20% d'entre eux ne guérissent pas (nous multiplions la probabilité d'appartenir à la sous-population qui peut recevoir le médicament avec la probabilité de ne pas guérir si on reçoit ce médicament). Vous pouvez aussi tracer un arbre de probabilité pour résoudre ce problème, si vous préférez"),
  fail_if(~ TRUE, "Votre calcul ne me semble pas correct. Combien ne peuvent recevoir le médicament, combien n'y réagissent pas positivement ? C'est la clé.")
)
```

Un médecin a trois patients atteints de la maladie en question dans sa salle d'attente ((ils ne se connaissent pas et ne sont pas de la même famille). Quelle est la probabilité que ce médecin puisse soigner l'ensemble de ses patients à l'aide du nouveau médicament ?

```{r medic2_h2, exercise=TRUE}

```

```{r medic2_h2-hint-1}
# Connaissant la probabilité de ne pas être soigné calculée ci-dessus,
# il est facile d'en déduire la probabilité de guérir pour un patient.
# Appliquez ensuite cette probabilité pour les 3 patients simultanément.

 #### ATTENTION: Hint suivant = solution !####
```

```{r medic2_h2-solution}
0.4 * 0.4 * 0.4
```

```{r medic2_h2-check}
grade_result(
  pass_if(~ identical(.result, 0.4 * 0.4 * 0.4), "Vous avez parfaitement compris le principe. Pr{guérison} = 1 - Pr{non guérison} pour un patient. Comme ce sont des évènements indépendants pour les trois patients, nous devons ensuite multiplier les probabilités de chaque patient entre elles."),
  fail_if(~ TRUE, "Votre calcul n'est pas le bon. Commencez par déterminer la probabilité pour un patient. Ensuite appliquez-là pour les trois ensemble.")
)
```

## Conclusion

Bravo ! Vous venez de vérifier votre bonne compréhension des probabilités et des calculs de probabilités. Si certaines questions vous semblaient plus difficiles, revoyez la théorie correspondante avant de progresser dans la matière.

```{r comm_noscore, echo=FALSE}
question_text(
  "Laissez-nous vos impressions sur ce learnr",
  answer("", TRUE, message = "Pas de commentaires... C'est bien aussi."),
  incorrect = "Vos commentaires sont enregistrés.",
  placeholder = "Entrez vos commentaires ici...",
  allow_retry = TRUE
)
```
