---
title: "Corrélation"
author: "Guyliann Engels & Philippe Grosjean"
description: "**SDD I Module 6** Coefficients et tests de corrélation."
output:
  learnr::tutorial
tutorial:
  id: "A06Lc_correlation"
  version: 3.0.0/5
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
BioDataScience1::learnr_setup()
SciViews::R("infer", lang = "fr")
# Required for RSConnect
# SciViews::R
library(rlang)
library(data.table)
library(ggplot2)
library(tibble)
library(tidyr)
library(dplyr)
library(dtplyr)
library(broom)
library(forcats)
library(collapse)
library(fs)
library(data.trame)
library(svFast)
library(svTidy)
library(svMisc)
library(svBase)
library(svFlow)
library(data.io)
library(chart)
library(tabularise)
library(SciViews)
# ... more
library(readxl)
library(testthat)
library(equatags)
library(distributional)
library(inferit)
library(BioDataScience)
library(BioDataScience1)
```

```{r, echo=FALSE}
BioDataScience1::learnr_banner()
```

```{r, context="server"}
if (Sys.info()["user"] == "rstudio-connect") {
  options(learnr_user_id = session$user)
  assignInNamespace("default_user_id", function() {
    getOption("learnr_user_id", unname(Sys.info()["user"]))
  }, ns = "learnr")
}
BioDataScience1::learnr_server(input, output, session)
```

------------------------------------------------------------------------

## Objectifs

L'association entre deux variables peut être quantifiée grâce à différents descripteurs ou indices comme la covariance, l'indice de corrélation de Pearson, l'indice de corrélation de Spearman ou de Kendall.

On définit trois niveaux d'association impliquant des hypothèses de plus en plus fortes sur les mécanismes sous-jacents liant ces deux variables : **corrélation**, **relation** et **causalité**.

La corrélation est l'association de deux variables sans élaborer plus d'hypothèses sur les mécanismes éventuellement responsables de cette association.

Dans ce tutoriel, vous allez pouvoir autoévaluer votre capacité à :

-   comprendre les différents niveaux d'association de deux variables numériques

-   réaliser et interpréter un test de corrélation

## Biométrie de trois espèces d'iris

L'étude porte sur les variations morphologiques des fleurs de trois espèces d'iris qui sont *Iris setosa* Pall. ex Link (1820) , *Iris versicolor* L. (1753) et *Iris virginica* L. Ces espèces appartiennent à la famille des *Iridaceae*. Ces trois espèces se développent dans des milieux humides.

Edgar Anderson a récolté des mesures sur les fleurs des trois espèces d'iris en Gaspésie (Québec, Canada) en 1935. La péninsule de Gaspésie est entourée par l'estuaire du fleuve Saint-Laurent, le golfe du Saint-Laurent et la baie des Chaleurs.

Cinquante plants par espèce sont prélevés afin d'étudier leurs fleurs. Les longueurs et largeurs des pétales et des sépales sont mesurées (en centimètre).

![](images/iris.jpg)

Vous avez à votre disposition un jeu de données sur des mesures biométriques d'`iris`.

```{r, echo=TRUE}
# Importation des données
(iris <- read("iris", package = "datasets", lang = "fr"))
```

La fonction `skim()` du package {skimr} est très intéressante dans un carnet de notes pour découvrir les données. Avec une seule fonction, on se fait rapidement une première idée sur le nombre d'observations, le type des variables, la présence de valeurs manquantes...

```{r, echo=TRUE}
skimr::skim(iris)
```

Le jeu de données `iris` est parfaitement balancé avec 50 individus par groupe. Il contient quatre variables numériques et une variable facteur non ordonnée. Il ne contient pas de valeurs manquantes.

```{r iris_prepa}
iris <- read("iris", package = "datasets", lang = "FR")
```

Avant de poursuivre votre exploration des données, réalisez par vous-même différents graphiques pour découvrir ce jeu de données dans la zone ci-dessous.

```{r scatterplot_noscore, exercise=TRUE, exercise.setup="iris_prepa"}
chart(data = iris, ___ ~ ___ %col=% species) +
  geom_point()
```

## Matrice de corrélation

Réalisez une matrice de corrélation afin de mettre en avant la corrélation linéaire entre l'ensemble des variables numériques disponibles. Effectuez ensuite un résumé de cette matrice.

<!-- Attention que le snippet se réfère à cor() alors qu'on utilise correlation() ici => homogénéiser, ou au moins le préciser. Ce serait aussi utile de rappeler les variables du tableau en début de cet exercice car on est passé à une autre page! -->

<!-- 💬 **Un snippet est disponible pour réaliser une matrice de corrélation à l'aide d'une fonction différente `.escor: correlation matrix`.** -->

```{r is_corr_h2, exercise=TRUE, exercise.lines=2, exercise.setup="iris_prepa"}
iris_corr <- correlation(___[, ___:___], use = "complete.obs", method = "___")
summary(iris_corr)
```

```{r is_corr_h2-hint-1}
iris_corr <- correlation(DF[, INDEX_EXPRESSION], use = "complete.obs", method = "pearson")
summary(iris_corr)
# La sélection se fait sur base des indices des colonnes en R de base.
# Vous ne devez sélectionner que les variables numériques.

 #### ATTENTION: Hint suivant = solution !####
```

```{r is_corr_h2-solution}
## Solution ##
iris_corr <- correlation(iris[, 1:4], use = "complete.obs", method = "pearson")
summary(iris_corr)
```

```{r is_corr_h2-check}
grade_code("Le résumé de la matrice de corrélation utilise un système de légende particulier pour mettre en avant les variables les plus corrélées. Assurez-vous de bien comprendre les symboles utilisés ., + ou encore B (explicité en dernière ligne).")
```

Réalisez un graphique de votre objet `iris_corr`. N'affichez que la partie inférieure de votre matrice. Tout comme le résumé de la matrice ne montre que la partie inférieure de la matrice, il n'est pas utile de représenter la matrice complète sous forme graphique, car la matrice est symétrique autour de sa diagonale.

```{r iris_corr}
iris <- read("iris", package = "datasets", lang = "FR")
iris_corr <- correlation(iris[, 1:4], use = "complete.obs", method = "pearson")
```

```{r is_plot_h2, exercise=TRUE, exercise.setup="iris_corr"}
___(iris_corr, type = "___")
```

```{r is_plot_h2-hint-1}
plot(iris_corr, type = "___")

#### ATTENTION: Hint suivant = solution !####
```

```{r is_plot_h2-solution}
## Solution ##
plot(iris_corr, type = "lower")
```

```{r is_plot_h2-check}
grade_code("Ce graphique permet d'avoir une information visuelle simple à analyser. La couleur utilisée, l'élongation de l'ellipse et le sens de l'ellipse sont porteurs d'information sur la corrélation entre les paires de variables étudiées. Répondez à la question suivante.")
```

```{r quiz_corr}
question("Quelles sont les variables les plus fortement corrélées (sélectionnez les deux variables les plus corrélées) ?",
  answer("sepal_length"),
  answer("sepal_width"),
  answer("petal_length", correct = TRUE),
  answer("petal_width", correct = TRUE),
  allow_retry = TRUE,
  incorrect = "Recommencez afin de trouver la bonne réponse.",
  correct = "Vous avez su décoder avec succès votre matrice de corrélation.")
```

## Test de corrélation

Suite à votre matrice de corrélation linéaire avec l'indice de corrélation de Pearson, vous avez pu observer que la longueur des pétales et la largeur des pétales ont un coefficient de corrélation supérieur à 0.9. Cependant, avoir une valeur élevée ne permet pas d'affirmer que la corrélation est **significative** entre ces deux variables. Utilisez un test de corrélation de Pearson pour répondre à la question suivante : **est-ce que la largeur des pétales (`petal_width`) et la longueur des pétales (`petal_length`) sont significativement corrélées au seuil alpha de 5% ?**

```{r cortest_h2, exercise=TRUE, exercise.lines=3, exercise.setup="iris_prepa"}
cor.test(data = ___, ~ ___ + ___,
  alternative = "two.sided", method = "___")
```

```{r cortest_h2-hint-1}
cor.test(data = iris, ~ ___ + ___,
  alternative = "two.sided", method = "___")

#### ATTENTION: Hint suivant = solution !####
```

```{r cortest_h2-solution}
## Solution ##
cor.test(data = iris, ~ petal_width + petal_length,
  alternative = "two.sided", method = "pearson")
```

```{r cortest_h2-check}
grade_code("Vous venez de réaliser un test de corrélation de Pearson. Il ne vous reste plus qu'à répondre à la question ci-dessous pour vous assurer que vous êtes capable d'interpréter ce test d'hypothèse.")
```

```{r quiz_test}
question("La largeur des pétales (`petal_width`) et la longueur des pétales (`petal_length`) sont-elles significativement corrélées au seuil alpha de 5% ?",
  answer("Oui", correct = TRUE),
  answer("Non"),
  correct = "Avec une valeur *p* inférieure au seuil alpha de 5%, il y a rejet de l'hypothèse nulle.",
  incorrect = "Il semble que vous n'avez pas su correctement interpréter ce test. Vérifez votre compréhension de ses hypothèses nulle et alternative, ainsi que la façon de décider laquelle favoriser.",
  allow_retry = TRUE)
```

## Conclusion

Bravo ! Vous venez de terminer le tutoriel sur la premier niveau d'association de paires de variables qu'est la corrélation. Au cours de science des données II, vous approfondirez ces notions en utilisant la modélisation pour représenter des **relations** entre variables.

```{r comm_noscore, echo=FALSE}
question_text(
  "Laissez-nous vos impressions sur cet outil pédagogique",
  answer("", TRUE, message = "Pas de commentaires... C'est bien aussi."),
  incorrect = "Vos commentaires sont enregistrés.",
  placeholder = "Entrez vos commentaires ici...",
  allow_retry = TRUE
)
```

```{css, echo=FALSE}
@media print {
  .topics {
    width: 100% !important;
    padding: 0 0.5em 0 !important;
  }
  .topicsContainer, .topicsContainer *, .learnrBanner, .learnrBanner *, .topicActions, .topicActions * {
    display: none !important;
  }
}
```
