---
title: "Corrélation"
author: "Guyliann Engels & Philippe Grosjean"
description: "**SDD I Module 12** Le coefficient de corrélation et le test de corrélation."
output:
  learnr::tutorial
tutorial:
  id: "A12La_correlation"
  version: 2.0.0/5
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
BioDataScience1::learnr_setup()
SciViews::R()
```

```{r, echo=FALSE}
BioDataScience1::learnr_banner()
```

```{r, context="server"}
BioDataScience1::learnr_server(input, output, session)
```

----

## Objectifs

L'association entre deux variables peut être quantifiée grâce à différentes descripteurs ou indices comme la covariance, l'indice de de corrélation de Pearson ou encore l'indice de de corrélation de Spearman. On définit 3 niveaux d'association impliquant des hypothèses de plus en plus fortes sur le mécanismes sou-jacent liant ces deux variables. On a la corrélation, la relation et la causalité. Les indices de corrélation et les tests de corrélation permettent comme leur nom l'indique de mettre un avant une corrélation entre deux variables. La corrélation est simplement l’association de deux variable sans élaborer plus d’explication sur le mécanisme sous-jacent qui lie ces variables.

Dans ce tutoriel, vous allez pouvoir auto-évaluer votre capacité à\ :

- Comprendre les différents niveaux d’association de deux variables numériques.

- Réaliser et interpréter un test de corrélation

## Biométrie de trois espèces d'Iris

L'étude porte sur les variations morphologiques des fleurs de trois espèces d'iris que sont *Iris setosa* Pall. ex Link (1820) , *Iris versicolor* L. (1753) et *Iris virginica* L.. Ces espèces appartiennent à la famille des Iridaceae. Ces trois espèces se développent dans des milieux humides. 

Edgar Anderson a récolté les trois espèces d'iris en Gaspésie (Québec, Canada) en 1935 au moment de la floraison. La péninsule de Gaspésie est entourée par l'estuaire du fleuve Saint-Laurent, le golfe du Saint-Laurent et la baie des chaleurs.

Cinquante plants par espèce sont prélevés afin d'étudier la biométrie de leurs fleurs. La longueur et la largeur des pétales et sépales sont mesurées (en centimètre).

![](images/iris.jpg)   

Vous avez à votre disposition un jeu de données sur des mesures biométriques d'iris.

```{r, echo = TRUE}
# Importation des données
(iris <- read("iris", package = "datasets", lang = "fr"))
```


La fonction `skim()` du package `skimr` est une fonction très intéressante dans un carnet de notes pour découvrir les données. Avec une seule fonction, on obtient une première idée sur le nombre d'observations, le type des variables, la présence de valeurs manquantes,...

```{r, echo=TRUE}
skimr::skim(iris)
```

Le jeu de données iris est parfaitement balancé avec 50 individus par groupe. Il comprends 4 variables numériques et 1 variable facteur non-ordonnée. Il ne comprend pas de valeurs manquantes.

```{r iris_prepa}
iris <- read("iris", package = "datasets", lang = "FR")
```

Avant de poursuivre votre exploration des données, réalisez différents graphiques en fonction des différents variables. 

```{r scatterplot_noscore, exercise=TRUE, exercise.setup = "iris_prepa"}
chart(data = iris, ___ ~ ___ %col=% species) +
  geom_point()
```

## Matrice de corrélation 

Réalisez une matrice de corrélation afin de mettre en avant la corrélation linéaire entre l'ensemble des variables numérique disponibles. Proposez ensuite un résumé de cette matrice.

💬 **Un snippet est disponible pour réaliser cet exercice `.escor: correlation matrix`.**

```{r is_corr_h2, exercise=TRUE, exercise.lines=1, exercise.setup="iris_prepa"}
iris_corr <- correlation(___[, ___:___], use = "complete.obs", method = "___")
summary(iris_corr)
```

```{r is_corr_h2-hint-1}
iris_corr <- correlation(DF[, INDEX_EXPRESSION], use = "complete.obs", method = "pearson")
summary(iris_corr)
# La sélection se fait sur base des indices des colonnes en R de base. 
# Vous ne devez sélectionner que les variables numériques.

 #### ATTENTION: Hint suivant = solution !####
```

```{r is_corr_h2-solution}
iris_corr <- correlation(iris[, 1:4], use = "complete.obs", method = "pearson")
summary(iris_corr)
```

```{r is_corr_h2-check}
grade_code("Le résume de la matrice de corrélation utilise un système de légende particulier pour mettre en avant les variables les plus corrélées. Assurez-vous de bien comprendre les symboles utilisés ., + ou encore B.")
```

Réalisez un graphique de votre objet `iris_corr`. N'affichez que la partie supérieur de votre matrice. Tout comme le résumé de la matrice ne montre que la partie inférieure de la matrice, il n'est pas utile de représenter la matrice complète sous forme graphique. La matrice est symétrique.

```{r iris_corr}
iris <- read("iris", package = "datasets", lang = "FR")
iris_corr <- correlation(iris[, 1:4], use = "complete.obs", method = "pearson")
```

```{r is_plot_h2, exercise=TRUE, exercise.lines=1, exercise.setup="iris_corr"}
___(iris_corr, type = "___")
```
 
```{r is_plot_h2-hint-1}
plot(iris_corr, type = "___")

#### ATTENTION: Hint suivant = solution !####
```
 
```{r is_plot_h2-solution}
plot(iris_corr, type = "upper")
```

```{r is_plot_h2-check}
grade_code("Ce graphique permet d'avoir une information visuelle simple à analyser. La couleur utilisée, l'élongation de l'éllipse et le sens de l'ellipse sont porteur d'information sur la corrélation entre les paires de variabls étudiées. Répondez à la question suivante.")
```

```{r quiz_corr}
question("Quelles sont les variables les plus fortement corrélées ?(sélectionnez les deux variables les plus corrélées)",
         answer("sepal_length"),
         answer("sepal_width"),
         answer("petal_length", correct = TRUE),
         answer("petal_width", correct = TRUE),
           allow_retry = TRUE,
           incorrect = "Recommencez afin de trouver la bonne réponse",
           correct = "Bravo, c'est correct. Vous avez su décoder avec succès votre matrice de corrélation")
```

## Test de corrélation 

Suite à votre matrice de corrélation linéaire avec l'indice de corrélation de Pearson, vous avez pu observer que la longueur des pétales et la largeur des pétales ont un indice de corrélation supérieur à 0.9. Cependant, avoir une valeur élevé ne permet pas d'affirmer qu'il y une corrélation significative entre ces deux variables. 

Utilisez un test de corrélation afin de répondre à la question suivante : **est ce que la largeur des pétales (petal_width) et la longueur des pétales (petal_length) sont significativement corrélées au seuil alpha de 5% ? **

```{r cortest_h2, exercise=TRUE, exercise.lines=3, exercise.setup="iris_prepa"}
cor.test(data = ___, ~ ___ + ___,
  alternative = "two.sided", method = "___")
```

```{r cortest_h2-hint-1}
cor.test(data = iris, ~ ___ + ___, 
  alternative = "two.sided", method = "___")

#### ATTENTION: Hint suivant = solution !####
```

```{r cortest_h2-solution}
cor.test(data = iris, ~ petal_width + petal_length, 
  alternative = "two.sided", method = "pearson")
```

```{r cortest_h2-check}
grade_code("Bravo ! Vous venez de réaliser un test de corrélation de Pearson. Il ne vous reste plus qu'à répondre à la question ci-dessous. Pour vous assurer que vous êtes capable de d'interpréter ce test d'hypothèse.")
```

```{r quiz_test}
question("Est ce que la largeur des pétales (petal_width) et la longueur des pétales (petal_length) sont significativement corrélées au seuil alpha de 5% ?",
  answer("Oui", correct = TRUE),
  answer("Non"),
  correct = "Bravo, c'est exact ! Avec une valeur inférieur au seuil alpha, il y a rejet de l'hypothèse nulle.", 
  incorrect = "Il semble que vous n'avez pas su déterminer correctement l'hypothèse nulle et l'hypothèse alternative.", 
  allow_retry = TRUE)
```

## Conclusion

Bravo! Vous venez de terminer le tutoriel sur la premier niveau d'association de paires de variables qu'est la corrélation.

```{r comm_noscore, echo=FALSE}
question_text(
  "Laissez-nous vos impressions sur cet outil pédagogique",
  answer("", TRUE, message = "Pas de commentaires... C'est bien aussi."),
  incorrect = "Vos commentaires sont enregistrés.",
  placeholder = "Entrez vos commentaires ici...",
  allow_retry = TRUE
)
```
