---
title: "ANOVA à 2 facteurs"
author: "Guyliann Engels & Philippe Grosjean"
description: "**SDD I Module 11** Analyse de variance à deux facteurs, différents modèles."
tutorial:
  id: "A11La_anova2"
  version: 2.0.0/11
output:
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
BioDataScience1::learnr_setup()
SciViews::R()
# data ----
set.seed(43)
production <- c(12, 13, 14, 11, 10, 13, 14, 12, 11, 10, 10, 12, 12, 13, 14, 10, 14, 13, 10, 14,
    17, 13, 17, 14, 13, 13, 17, 14, 17, 13, 10, 9, 10, 12, 11, 9, 10, 12, 10, 11,
    11, 9, 11, 11, 12, 10, 9, 13, 11, 10, 13, 12, 11, 15, 14, 11, 15, 13, 14, 12) 
production <- round(production + rnorm(n = length(production), sd = 0.5), 1)

milk_production <- tibble(
  ration = ordered(rep(c("low", "high"), each = 30), levels = c("low", "high")), 
  feed = factor(rep(rep(c("hay", "grass", "maize silage"), each = 10), times = 2)), 
  milk = production)

milk_production <- labelise(milk_production, 
  label = list(ration = "Ration", feed = "Type d'alimentation", milk = "Quantité de lait"),
  units = list(milk = "L/J"))
```

```{r, echo=FALSE}
BioDataScience1::learnr_banner()
```

```{r, context="server"}
BioDataScience1::learnr_server(input, output, session)
```

## Objectifs

Lors de la réalisation du tutoriel précédent portant sur l'ANOVA à un facteur, vous avez étudié une variable réponse pour différents niveaux d'une variable facteur. Dans le [module 11](https://wp.sciviews.org/sdd-umons/?iframe=wp.sciviews.org/sdd-umons-2020/variance2.html) du cours, vous avez vu qu'il était également possible de considérer deux variables facteurs avec une **ANOVA à deux facteurs**. 

Ce tutoriel vous permettra d'auto-évaluer votre aptitude à\ :

- présenter vos donnés avant une ANOVA à deux facteurs

- effectuer une ANOVA à deux facteurs 

- interpréter correctement les résultats obtenus

Assurez-vous d'avoir compris les sections [11.1](https://wp.sciviews.org/sdd-umons/?iframe=wp.sciviews.org/sdd-umons-2020/anova-%25C3%25A0-deux-facteurs.html), [11.2](https://wp.sciviews.org/sdd-umons/?iframe=wp.sciviews.org/sdd-umons-2020/mod%25C3%25A8le-sans-interactions.html) et [11.3](https://wp.sciviews.org/sdd-umons/?iframe=wp.sciviews.org/sdd-umons-2020/mod%25C3%25A8le-crois%25C3%25A9-complet.html) du cours. Vous devez également maitriser les différentes notions vue dans le [module 10](https://wp.sciviews.org/sdd-umons/?iframe=wp.sciviews.org/sdd-umons-2020/variance.html) relatives à l’ANOVA à un facteur.

## Rendement laitier

Lors de ce tutoriel, vous allez réaliser une analyse visant à étudier l'influence du type et de la quantité d'alimention sur le rendement laitier. 

![](images/rendement_laitier.jpg)

Le jeu de données `milk_production` comprend les variables suivantes\ :

- `ration`\ : la quantité d'alimentation apportée (low = ration faible ou high = ration élevée)
- `feed`\ : le type d'alimentation fourni (hay = foin, grass = herbe ou maize silage = maïs ensilés)
- `milk`\ : la quantité de lait en L/jour

```{r}
glimpse(milk_production)
```

La question biologique à laquelle vous allez essayer de répondre est la suivante : **y a t'il une différence de rendement dans la production de lait en fonction de la ration et du type d'alimentation donnée aux vaches ?**

### Description graphique des données

Différents graphiques vous permettent de visualiser correctement les observations de votre jeu de données avant la réalisation de votre ANOVA à deux facteurs. À partir du jeu de données (`milk_production`), réalisez un premier graphique représentant la moyenne du rendement de la production de lait (`milk`) en fonction du type d'alimentation (`feed`) et en utilisant la couleur pour la quantité d'alimentation apportée (`ration`). 

```{r lait_graph_h2, exercise = TRUE}
chart(___, ___ ~ ___ %col=% ___) +
  geom_point(alpha = 0.4, position = position_dodge2(0.4, padding = 0.1)) +
  stat_summary(geom = "point", fun = "mean", position = position_dodge(0.4), size = 4)
```

```{r lait_graph_h2-hint-1}
chart(milk_production, ___ ~ feed %col=% ___) +
  geom_point(alpha = 0.4, position = position_dodge2(0.4, padding = 0.1)) +
  stat_summary(geom = "point", fun = "mean", position = position_dodge(0.4), size = 4)

 #### ATTENTION: Hint suivant = solution !####
```

```{r lait_graph_h2-solution}
chart(milk_production, milk ~ feed %col=% ration) +
  geom_point(alpha = 0.4, position = position_dodge2(0.4, padding = 0.1)) +
  stat_summary(geom = "point", fun = "mean", position = position_dodge(0.4), size = 4)
```

```{r lait_graph_h2-check}
grade_code("Le graphique que vous venez de réalisez permet de comparer le rendement laitier en fonction de la ration par types d'alimentation. Le choix du nuage de point avec l'ajout de la moyenne est tout indiqué lorsque le nombre d'observations est faible. Les aliments ensilés permettent d'avoir mettent en avant une production plus importante de lait par jour.")
```

### Visualisation graphique des interactions 

Dans un modèle d'ANOVA à deux facteurs, il faut aussi considérer que l'effet d'un facteur sur la variable explicative peut-être différent suivant les modalités de l'autre facteur mis en jeu. On parle alors d'interactions. Il est possible de les explorer en réalisant un graphique des interactions qui vous permettra de visualiser les écarts des moyennes respectives des différentes sous-populations. 

Réalisez sur le jeu de données `milk_production` le graphique permettant de visualiser les interactions entre les variables facteurs (`feed` et `ration`). Commencez par calculer la moyenne du rendement par types d'alimentation et par ration. Réalisez ensuite un graphique du rendement moyen en fonction des rations par types d'alimentation en reliant les points des moyennes par des lignes pour chaque type d'alimentation. 


```{r interaction_h3, exercise = TRUE}
___ %>.%
  ___(., ___, ___) %>.% 
  ___(., milk_mean = ___(___)) %>.%  # moyenne du rendement 
  chart(data = ___, ___ ~ ___ %col=% ___ %group=% ___) + # graphique du rendement moyen en fonction de la quantité d'aliment par type
  ___() + # lignes
  ___() # points
```

```{r interaction_h3-hint-1}
DF %>.%
  group_by(., FACTOR1, FACTOR2) %>.% # regroupement par type d'alimentation et par ration
  summarise(., milk_mean = mean(VARNUM)) %>.%   # moyenne du rendement 
  chart(data = ___, ___ ~ ___ %col=% ___ %group=% ___) + # graphique du rendement moyen en fonction de la quantité d'aliment par type
  ___() + # lignes
  # points
  ___()
```

```{r interaction_h3-hint-2}
milk_production %>.%
  group_by(., feed, ration) %>.% # regroupement par type et quantité d'alimentation
  summarise(., milk_mean = mean(milk)) %>.% # moyenne du rendement 
  chart(data = ., VARNUM ~ FACTOR1 %col=% FACTOR2 %group=% FACTOR2) +
  ___() + # lignes
  ___() # points

 #### ATTENTION: Hint suivant = solution !####
```

```{r interaction_h3-solution}
milk_production %>.%
  group_by(., feed, ration) %>.% # regroupement par type et quantité d'alimentation
  summarise(., milk_mean = mean(milk)) %>.%  # moyenne du rendement 
  chart(data = ., milk_mean ~ feed %col=% ration %group=% ration) + 
  geom_line() + # lignes
  geom_point() # points
```

```{r interaction_h3-check}
grade_code("Ce graphique permet d'analyser les interactions entre les prédicteurs (rations et types d'alimentation) et la réponse (rendement laitier). Lorsque les lignes sont parallèles, il n'y a pas d'interaction à l'inverse, lorsque les lignes ne sont pas parallèles, il existe probablement une interaction. Moins les lignes sont parallèles plus l'interaction sera forte. Dans notre cas, nous serions tenté de penser qu'il n'y a pas d'intéractions. Cependant, il reste plus sur de réaliser un modèle croisé complet dans un premier temps. Puis de réaliser un modèle croisé sans intéractions si cela se justifie.")
```

### Description numérique sous la forme d'un tableau

Vous pouvez également résumer vos données sous la forme d'un tableau. À partir du jeu de données `milk_production`, réalisez un tableau de résumé contenant la moyenne et  l'écart-type du rendement de la production de lait (`milk`) par type d'alimentation (`feed`) et par quantité des rations (`ration`). Ajoutez également le nombre d'observation par sous-groupe. 

💬 **Le code ci-dessous correspond au snippet `.hmanova2desc`.**

```{r lait_tab_h2, exercise=TRUE}
___ %>.%
  group_by(., ___, ___) %>.%
  summarise(., mean = mean(___), sd = sd(___), count = sum(!is.na(___)))
```

```{r lait_tab_h2-hint-1}
DF %>.%
  group_by(., XFACTOR1, XFACTOR2) %>.%
  summarise(., mean = mean(YNUM), sd = sd(YNUM), count = sum(!is.na(YNUM)))

 #### ATTENTION: Hint suivant = solution !####
```

```{r lait_tab_h2-solution}
milk_production %>.%
  group_by(., feed, ration) %>.%
  summarise(., mean = mean(milk), sd = sd(milk), count = sum(!is.na(milk)))
```

```{r lait_tab_h2-check} 
grade_code("Le tableau que vous avez obtenu est parfait pour présenter vos résultats avant de réaliser une ANOVA à deux facteurs. Attention, ce tableau est en anglais. N'oubliez pas de le traduire si votre rapport est en français.")
```

Les graphiques et la tableau résumant des données que vous venez de réaliser vous ont permis de prendre connaissance des données sur lesquelles vous travaillez. Vous allez maintenant pouvoir passer à la réalisation de votre test d'hypothèse. Vous commencez maintenant à avoir l'habitude, avant de pouvoir le réaliser, vous allez devoir vérifier si les conditions d'application du test sont respectées. 

### Vérification des conditions d'applications

Si vous souhaitez réaliser une ANOVA à deux facteurs, vous devez respecter les conditions d'application suivantes\ :

- échantillon représentatif,
- observations indépendantes,
- une variable réponse quantitative,
- deux variables explicatives qualitatives à deux niveaux ou plus,
- distribution **normale** des résidus,
- **homoscédasticité** (même variance intragroupe).

Pour vérifier si l'échantillonage est représentatif de la population et si les observations sont bien indépendantes les unes des autres, vous pouvez vous référer au protocole expérimental. Dans le cas de cette étude, les individus ont bien été sélectionnés par un processus aléatoire. L'attribution des niveaux de la ration et du type d'alimentation a également été réalisée de manière aléatoire. 

Le type de variable peut se vérifier en consultant les données. Attention de bien vous assurer que les deux variables explicatives que vous utilisez sont bien de type facteur non ordonné (`factor`) ou facteur ordonné (`ordered`) dans R.  

Vous allez maintenant pouvoir vérifier les deux dernières conditions sur votre échantillon. Les données par sous-population étant en nombre suffisant, vous allez pouvoir appliquer un test permetant de vérifier l'homoscédasticité des variances. Vous devrez aussi vérifier que la distribution des résidus suit bien une distribution Normale. Cette étape se fait après la réalisation du test de l'ANOVA à deux facteurs mais avant son interprétation. 

### Test de Bartlett

Pour évaluer l'homoscédasticité de vos données, vous allez utiliser le test de Bartlett. Afin d'interpréter correctement le test que vous allez réaliser, commencez par poser les hypothèses de départ. Le test de Barlett doit se faire sur toutes les sous-populations. Pour cela, vous pouvez utiliser la fonction `interaction()` en lui renseignant les variables explicatives. Cette fonction, qui calculera pour vous tous les sous-groupe possible, peut directement être utilsée dans la formule. 

Réalisez un test de bartlett avec interaction sur les données du jeu `milk_production` de la quantité de lait (`milk`) en fonction du type d'alimentation (`feed`) et de la ration (`ration`) .

💬 **Le code ci-dessous correspond au snippet `.hvbartlett`.**

```{r lait_bart_h2, exercise = TRUE}
bartlett.test(data = ___, ___ ~ ___(___, ___))
```

```{r lait_bart_h2-hint-1}
bartlett.test(data = DF, YNUM ~ interaction(XFACTOR1, XFACTOR2))

 #### ATTENTION: Hint suivant = solution !####
```

```{r lait_bart_h2-solution}
bartlett.test(data = milk_production, milk ~ interaction(feed, ration))
```

```{r lait_bart_h2-check}
grade_code("Vous avez correctement réaliser l'instruction R pour ce test. Saurez-vous l'interpréter correctement en répondant à la question suivante ? ")
```

```{r lait_bart_quiz}
question("Au seuil alpha de 5%, y a t'il homoscédasticité ?",
         answer("Oui", correct = TRUE),
         answer("Non"),
         allow_retry = TRUE,
         random_answer_order = TRUE,
         incorrect = "Avez-vous bien défini les hypothèses H0 et H1 de votre test de Bartlett ?",
         correct = "Bravo, c'est correct. La valeur de p que vous avez obtenue étant plus grande que le seuil alpha de 5%, vous ne rejettez donc pas H0 et pouvez considérer qu'il y a homoscédasticité.")
```

## Anova à 2 facteurs

### Anova à 2 facteurs avec interactions

Vous allez maintenant pouvoir réaliser une ANOVA à 2 facteurs avec interaction.

Comme pour l'ANOVA à un facteur, vous allez utiliser la fonction `lm()` qui demande un jeu de données et une formule avant de lui appliquer la méthode `anova()` afin de produire le tableau de l'analyse de la variance.

Le modèle croisé avec interactions peut s'écrire de 2 manières dans R.

```{r, echo=TRUE, eval=FALSE}
anova(anova. <- lm(data = DF, YNUM ~ XFACTOR1 + XFACTOR2 + XFACTOR1:XFACTOR2))
```

où `XFACTOR1:XFACTOR2` représente le terme correspondant à l'interaction. 

Ce modèle peut aussi se simplifier en utilisant `*` à la place du `+` entre les deux variables facteurs ce qui signifiera implicitement de tenir compte des interactions. Cela donne, 

```{r, echo=TRUE, eval=FALSE}
anova(anova. <- lm(data = DF, YNUM ~ XFACTOR1 * XFACTOR2))
```

Réalisez votre modèle ANOVA à deux facteurs avec interactions sur le rendement de lait en fonction du type d'alimentation et de la ration.

💬 **Le code ci-dessous correspond au snippet `.hmanova2`**.

```{r lait_interaction_h2, exercise = TRUE}
anova(anova. <- lm(data = ___, ___ ~ ___ * ___))
```

```{r lait_interaction_h2-hint-1}
anova(anova. <- lm(data = milk_production, ___ ~ ___ * ___))

 #### ATTENTION: Hint suivant = solution !####
```

```{r lait_interaction_h2-solution}
anova(anova. <- lm(data = milk_production, milk ~ feed * ration))
```

```{r lait_interaction_h2-check}
grade_code("Bravo, vous avez encodé correctement votre modèle. N'oubliez pas qu'avant de pouvoir interpréter vos résultats vous devez vérifier la normalité des résidus. Dans cette analyse, les résidus suivent une distribution normale, les conditions de votre ANOVA à deux facteurs sont donc remplies. Vous pouvez interpréter vos résultats en répondant à la question ci-dessous.")
```

```{r lait_interaction_quiz}
question("Y a t'il une intéraction significative entre les rations et les types d'alimention au seuil alpha de 5% ?",
     answer("oui"),
     answer("non", correct = TRUE),
     allow_retry = TRUE,
     incorrect = "L'interprétation que vous avez faite n'est pas la bonne. Retournez voir dans le cours comment se prend la décision du rejet ou non de H0 en comparant la valeur de p obtenue au seuil alpha 5%.",
     correct = "Bravo, c'est correct. Il n'y a pas d'interactions entre les 2 facteurs. Avant de continuez l'analyse, vous allez simplifier le modèle en retirant le terme correspondant à l'interaction.")
```

### Anova à 2 facteurs sans interactions

Lorsque vous n'avez pas d'interactions, vous pouvez simplifier votre modèle en supprimant le terme correspondant à l'interaction. Votre modèle s'écrit donc de cette manière. 

```{r, echo=TRUE, eval=FALSE}
anova(anova. <- lm(data = DF, YNUM ~ XFACTOR1 + XFACTOR2))
```

Réalisez votre modèle ANOVA à deux facteurs sans interactions sur le rendement de lait en fonction de la ration et du type d'alimentation.

💬 **Le code ci-dessous correspond au snippet `.hmanova2noint`**.

```{r lait_croise_h2, exercise = TRUE}
anova(anova. <- lm(data = ___, ___ ~ ___ + ___))
```

```{r lait_croise_h2-hint-1}
anova(anova. <- lm(data = lait, ___ ~ ___ + ___))

 #### ATTENTION: Hint suivant = solution !####
```

```{r lait_croise_h2-solution}
anova(anova. <- lm(data = milk_production, milk ~ feed + ration))
```

```{r lait_croise_h2-check}
grade_code("Bravo, vous avez réussi à simplifier votre modèle. Dans cette analyse, les résidus suivent également une distribution normale, les conditions de votre ANOVA à deux facteurs sont donc également remplies. Vous pouvez interpréter vos résultats en répondant aux questions ci-dessous.")
```

```{r lait_croise_quiz}
quiz(question("Y a t'il une différence significative entre les deux rations au seuil alpha de 5% ?",
     answer("oui", correct = TRUE),
     answer("non"),
     allow_retry = TRUE,
     incorrect = "L'interprétation que vous avez faite n'est pas la bonne. Retournez voir dans le cours comment se prend la décision du rejet ou non de H0 en comparant la valeur de p obtenue au seuil alpha 5%.",
     correct = "Votre interprétation est correct, il y a bien une différence significative, au seuil alpha de 5%, entre les rations données sur le rendement de lait."),
     question("Y a t'il une différence significative entre les trois types d'alimentation au seuil alpha de 5% ?",
     answer("oui", correct = TRUE),
     answer("non"),
     allow_retry = TRUE,
     incorrect = "L'interprétation que vous avez faite n'est pas la bonne. Retournez voir dans le cours comment se prend la décision du rejet ou non de H0 en comparant la valeur de p obtenue au seuil alpha 5%.",
     correct = "Votre interprétation est correct, il y a bien une différence significative, au seuil alpha de 5%, entre les types d'alimentation donnés sur le rendement de lait.")
     )
```

## Tests post-hoc

Votre ANOVA à 2 facteurs sans interactions indique qu'il y a une différence significative entre les rations données et qu'il y a au moins une des moyennes qui est significativement différente des autres au seuil $\alpha$ de 5% pour le type d'alimentation. Un test post-hoc de comparaisons multiples vous permettra d'avoir plus de précision. 

💬 **Ce code correspond au snippet `.hmanovamult`**

```{r lait_prepa}
anova. <- lm(data = milk_production, milk ~ feed + ration)
```

```{r lait_post_h2, exercise = TRUE, exercise.setup = "lait_prepa"}
summary(anovaComp. <- confint(multcomp::glht(anova.,
  linfct = multcomp::mcp(___ = "Tukey", ___ = "Tukey")))) # Add a second factor if you want
.oma <- par(oma = c(0, 5.1, 0, 0)); plot(anovaComp.); par(.oma); rm(.oma)
```

```{r lait_post_h2-hint-1}
summary(anovaComp. <- confint(multcomp::glht(anova.,
  linfct = multcomp::mcp(XFACTOR1 = "Tukey", XFACTOR2 = "Tukey")))) # Add a second factor if you want
.oma <- par(oma = c(0, 5.1, 0, 0)); plot(anovaComp.); par(.oma); rm(.oma)

 #### ATTENTION: Hint suivant = solution !####
```

```{r lait_post_h2-solution}
summary(anovaComp. <- confint(multcomp::glht(anova.,
  linfct = multcomp::mcp(feed = "Tukey", ration = "Tukey")))) # Add a second factor if you want
.oma <- par(oma = c(0, 5.1, 0, 0)); plot(anovaComp.); par(.oma); rm(.oma)
```

```{r lait_post_h2-check}
grade_code("Le test de comparaison n'est pas vraiment nécessaire dans le cas de la variable ration puisque vous n'avez que 2 niveaux. Mais cela vous permet de voir comment procéder lorsque le nombre de niveaux est plus important. Par contre, il prend tout son sens dans le cas de la variable alimentation pour identifier les paires de moyennes qui sont différentes. Identifiez les niveaux qui sont significativement différents en répondant à la question suivante.")
```

```{r lait_post_quiz}
question("Quels sont les niveaux significativement différents au seuil alpha de 5% ?",
         answer("foin par rapport à aliments ensilés", correct = TRUE),
         answer("herbe par rapport à aliments ensilés", correct = TRUE),
         answer("herbe par rapport à foin", correct = FALSE), 
  allow_retry = TRUE, correct = "Très bien ! En effet, la moyennes des aliments ensilés est différentes de celle du foin ou de l'herbe.",
  incorrect = "Votre réponse est peut-être incomplète ou vous avez commis une erreur dans l'interprétation. Regardez bien le tableau et le graphique issus de l'analyse post hoc.")
```

## Conclusion

Félicitation ! Vous avez réalisé une ANOVA à deux facteurs complète. N'oubliez pas de bien vérifier vos conditions d'application de test avant d'interpréter vos résultats. N'hésitez pas à simplifier votre modèle lorsque cela s'avère nécessaire.


```{r comm_noscore, echo=FALSE}
question_text(
  "Laissez-nous vos impressions sur cet outil pédagogique",
  answer("", TRUE, message = "Pas de commentaires... C'est bien aussi."),
  incorrect = "Vos commentaires sont enregistrés.",
  placeholder = "Entrez vos commentaires ici...",
  allow_retry = TRUE
)
```
